---
title: '1 Введение в основы'
description: "В этой главе вы сделаете первые шаги с R. Вы узнаете, как использовать консоль в качестве калькулятора и как назначать переменные. Вы также познакомитесь с базовыми типами данных в R. Давайте начнем!\n\n"
---

## Как это работает

```yaml
type: NormalExercise 
lang: r
xp: 100 
skills: 1
key: 1728f156db   
```


В редакторе справа вы должны ввести R-код для решения упражнений. Когда вы нажимаете кнопку «Отправить ответ», каждая строка кода интерпретируется и выполняется R, и вы получаете сообщение о том, был ли ваш код правильным. Вывод вашего R-кода отображается в консоли в нижнем правом углу.

R использует "#" знак для добавления комментариев, чтобы вы и другие могли понять, что такое код R. Так же, как Twitter! Комментарии не запускаются как R-код, поэтому они не будут влиять на ваш результат. Например, *Рассчитать 3 + 4* в редакторе справа - комментарий.

Вы также можете выполнять команды R прямо в консоли. Это хороший способ экспериментировать с R-кодом, так как ваше представление не проверено на правильность.


`@instructions`
- В редакторе справа уже есть пример кода. Вы можете видеть, какие строки являются реальным кодом R и какие комментарии?
- Добавьте строку кода, которая вычисляет сумму 6 и 12, и нажмите кнопку «Отправить ответ».

`@hint`
- Просто добавьте строку кода R, которая вычисляет сумму 6 и 12, как пример в примере кода!

`@pre_exercise_code`

```{r}
# Load datasets and packages here.
```


`@sample_code`

```{r}
# рассчитать 3 + 4


# рассчитать 6 + 12


```


`@solution`

```{r}
# рассчитать 3 + 4
3 + 4

# рассчитать 6 + 12
6 + 12 
```


`@sct`

```{r}
# Update this to something more informative.
success_msg("Some praise! Then reinforce a learning objective from the exercise.")
```


